# Homework 2
* CSCI-4930 ML Spring 2026
* Instructor: Dr. B

## Your name
* My name is  Kritika Joshi

## Rules to work with this notebook
* Make sure you have a `python 3.12.x` (e.g., Python 3.12.6) virtual environment ready for this homework.
* You can create new cells in this notebook.
* PLEASE DO NOT DELETE ANY EXISTING CELL FROM THIS NOTEBOOK.
* In a cell, you are allowed only to write codes/snippets right after the marking like `##YOUR CODE HERE`. You are not allowed to write anywhere else in that particular cell.
* Your work for each task has to be in the specific cells with the `##YOUR CODE HERE` marks, otherwise your submission won't be graded. Please do not delete the marks.
* After the ##YOUR CODE HERE marks, you may see command `raise NotImplementedError()` exception that alerts us that you did not attempt a particular task. Be sure to delete that command if you attempted to solve the task.

```python
#Please make sure you are using a python3.12.x interpreter 
# and install the packages from requirements.txt (provided)
#!pip install -r requirements.txt
```

## The imports

```python
#Some basic imports. You can import any package at any cell
import numpy as np
import pandas as pd

from sklearn.preprocessing import OneHotEncoder
from sklearn.preprocessing import OrdinalEncoder
from sklearn.preprocessing import StandardScaler


from sklearn.impute import SimpleImputer
from sklearn.pipeline import Pipeline


from sklearn.model_selection import train_test_split

from numpy.linalg import inv
from time import perf_counter
import math
import sys
import random
from tqdm import tqdm #good looking progress bars
```

```python
### Let's load the given dataset
# First load the dataset into pandas dataframe
full_dataset = pd.read_csv('dataset/baby-weights-dataset-2.csv',delimiter=',')
judge_dataset = pd.read_csv('dataset/judge-without-labels-2.csv',delimiter=',')
```

## TASK 1: 
Separate the full_dataset into two parts: X and y, where X denotes the input matrix containing only the input (i.e., independent explanatory) variables, and y denotes the target variable containing only the target values for exactly the same number of samples in the given full_dataset. 

<!-- grade_id: cell-task-1-ans, points: 0, grade: false, solution: true -->
```python
#Make sure you write your solution to this task below in this cell only.
X = pd.DataFrame() 
y = pd.DataFrame() 

# YOUR CODE HERE
X = full_dataset.drop(columns=['BWEIGHT'])
y = full_dataset['BWEIGHT']

#raise NotImplementedError("Task 1 not completed.")
```

<!-- grade_id: cell-task-1-test, points: 10, grade: true, solution: false -->
```python
#Let me test your code above
assert X.shape==(101400,36), "Shape of X does not match the expected shape."
assert y.shape==(101400,), "Shape of y does not match the expected shape."
print(f'Task 1 test complete. All good.')
```

Output:
```
Task 1 test complete. All good.
```

## TASK 2:
Compute mean, stdev, min, max, 25% percentile, median and 75% percentile of BWEIGHT target variable (i.e, the target y) in the full dataset (i.e., `y`), and print the computed values as a numpy array, `computed_values` containing these 7 results (respectively).


<!-- grade_id: cell-task-3-ans, points: 0, grade: false, solution: true -->
```python
#Make sure you write your solution to this task below in this cell only.
mean_val = 0
stdev_val = 0
min_val = 0
max_val = 0
percentile25_val = 0
median_val = 0
percentile75_val = 0

computed_values = np.array([])

# YOUR CODE HERE
mean_val = y.mean() 
stdev_val = y.std() 
min_val = y.min() 
max_val = y.max() 
percentile25_val = y.quantile(0.25) 
median_val = y.median()
percentile75_val = y.quantile(0.75)

computed_values = np.array([mean_val, stdev_val, min_val, max_val, 
                            percentile25_val, median_val, percentile75_val])

#raise NotImplementedError("Task 2 not completed.")
```

<!-- grade_id: cell-task-3-test, points: 10, grade: true, solution: false -->
```python
assert abs(sum([mean_val,stdev_val, min_val,max_val, percentile25_val,median_val,percentile75_val])-43.90001996) < 1e-4, "Your computed values do not contain correct answer(s)."
print(f'Task 2 test completed. All good.')
```

Output:
```
Task 2 test completed. All good.
```

## Task 3:
* Now that we have the dataset loaded, understood into our current workspace, according to the Machine Learning Pipeline we learned in the first class, the dataset needs to go through a rigorous `preprocessing` pipeline before we send it to train any machine learning model. 
* For the preprocessing of the dataet, we will use `sklearn.pipeline.Pipeline` and `sklearn.compose.ColumnTransformer` which ensures the following:
  * All preprocessing steps are learned only from the `training data` split, not the whole dataset. Remember, we must set-aside a split (`test split`) for an unbiased evaluation of our model later on. This way, `data leakage` can be prevented.
  * The same learned preprocessing transformations are applied consistently to both training and test splits.

* **Our preprocessing steps on the loaded dataset:**
  * Step 1: `Train-Test split`
  * Step 2: `Column Grouping`
  * Step 3: Prepare `Numeric` columns preprocessing pipeline
  * Step 4: Prepare `Ordinal` columns Preprocessing pipeline
  * Step 5: Prepare `Nominal` columns preprocessing pipeline
  * Step 6: Prepare `Default` columns preprocessing pipeline (for binary columns)
  * Step 7: Prepare `ColumnTransformer` combining three pipelines (steps 3,4,5)
  * Step 8: Train/fit the pipeline using the training split (`X_train`)
  * Step 9: Transform the training split (`X_train`) with the fitted transformer.
  * Step 10: Transform the test split with the same fitted/trained transformer
* Alright, let's start building preprocessing pipeline.

### Step 1: `Train-Test split` -- 
   - Splits the dataset into: Training set and Test set. Training set is used to learn preprocessing parameters and model weights, while the Test set is used only for evaluation. 
   - Given X representing the input matrix from the full_dataset, y being the target vector (the rightmost column of the full_dataset), obtained from Task 1, randomly split the (X,y) dataset into 75% for training and 25% for testing using the library function from the library [sklearn.model_selection](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html) . Please pass to the train_test_split function an additional argument random_state=45931.
   - Store the 4 splits as `X_train`, `X_test`, `y_train`, `y_test` respectively.
   - Save the ID column for `X_train` and `X_test` into `ID_train` and `ID_test` as list variable.
   - Now, drop the ID columns from both X_train and X_test

```python
#Step 1: Train-test split

#You are encouraged to use the `train_test_split()` method from scikit-learn
from sklearn.model_selection import train_test_split

X_train = pd.DataFrame()
X_test = pd.DataFrame()
y_train = pd.DataFrame()
y_test = pd.DataFrame()
ID_train = pd.DataFrame()
ID_test = pd.DataFrame()

# YOUR CODE HERE
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.25, random_state=45931)

ID_train = X_train[['ID']]
ID_test = X_test[['ID']]

X_train = X_train.drop(columns=['ID'])
X_test = X_test.drop(columns=['ID'])


#raise NotImplementedError("Task 3, Step 1 not completed.")
```

```python
#Let me test your code above
assert (X_train.shape, X_test.shape, y_train.shape, y_test.shape) == ((76050, 35), (25350, 35), (76050,), (25350,)), "Shapes of the training and test splits are not correct."
assert (ID_train.shape, ID_test.shape) == ((76050,1), (25350,1)), "Shapes of the ID_train and ID_test variables are not correct."
print(f'Task 3, Step 1 test complete. All good.')
```

Output:
```
Task 3, Step 1 test complete. All good.
```

### Step 2: `Column Grouping` --

| Type    | Examples        | Why separate?     |
| ------- | --------------- | ----------------- |
| Ordinal | FEDUCATION, MEDUCATION | Categorical values having natural order |
| Nominal | RACEMOM, RACEDAD, HISPMOM,HISPDAD, SEX, LOUTCOME  | Categorical values with no inherent order |
| Numeric | Remaining columns (except binary columns)     | Numeric values that needs scaling      |
| Binary  | Columns which are binary | Having values 0 or 1 |

Different data types require different preprocessing strategies.

```python
#Step 2: Let's do Column Grouping by listing the column names in respective list variables below.


ordinal_features = [] #list of ordinal column names
the_orders = {} #order of categories within each ordinal column as (key:value) pair.
                #Each ordinal column name will be the key,
		  # and list of categories denoting the order will be the value.
		  # For example: the_orders = { 'my_column_p':['val1', 'val2', 'val3'],
		  #                             'my_column_q':['cat8', 'cat9', 'cat10'],
                #              }
		  #              where, `my_column_p` and `my_column_q` are two ordinal columns. And, `val1 < val2 < val3` is the order of categories for my_column_p column, while cat8 < cat9 < cat10 is the order of categories for my_column_q column.

nominal_features = [] #list of nominal / categorical column names.


numeric_features = [] #list of numerical column names (except binary columns)

all_other_features = [] #list of all other features (the binary columns)

# YOUR CODE HERE
# Step 2: Column Grouping
# YOUR CODE HERE
ordinal_features = ['FEDUCATION', 'MEDUCATION']
the_orders = {
    'FEDUCATION': [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17],
    'MEDUCATION': [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17]
}
nominal_features = ['RACEMOM', 'RACEDAD', 'HISPMOM', 'HISPDAD', 'SEX', 'LOUTCOME']

# Using MAGE instead of MAGER, and other names from your printed list
numeric_features = ['MAGE', 'FAGE', 'GAINED', 'VISITS', 'WEEKS', 'TOTALP', 'CIGNUM', 'DRINKNUM']

# This catches the binary flags (ANEMIA, CARDIAC, etc.)
all_other_features = [col for col in X_train.columns if col not in (ordinal_features + nominal_features + numeric_features)]
```

```python
#Let me test your code above
assert len(ordinal_features) > 0, "There is at least one ordinal feature in the dataset. Please look for them."
assert len(the_orders) > 0, "There is at least one ordinal feature, and thus `the_order` dictionary must have at least one (key,value) pair. Please fix it."
assert len(nominal_features) > 0, "There is at least one nominal feature in the dataset. Please look for them."
assert len(numeric_features) > 0, "There is at least one numeric feature in the dataset. Please look for them."
assert len(all_other_features) > 0, "There is at least one feature not included in other 3 lists."
print(f'Task 3, Step 2 test complete. All good.')
```

Output:
```
Task 3, Step 2 test complete. All good.
```

### Step 3: Prepare `Numeric` columns preprocessing Pipeline
#### Step 3.0 Construct a `Pipeline` object named `numeric_transformer`
* A sklearn Pipeline allows you to sequentially apply a list of transformers to preprocess the data. 
* * For this `numeric_transformer` you are going to apply the following two transformations: i) Missing value imputation step, ii) Standard Scaling step.
```python
from sklearn.pipeline import Pipeline
numeric_transformer = Pipeline(steps=[
    #list of steps as a list of tuples.
    #(name_of_step1, processor_object_1),
    #(name_of_step2, processor_object_2),
    # ...
])
```

#### Step 3.1 `Missing Value Imputation of the Numeric features` --
   - Replaces missing numeric values with a meaningful value (e.g., the median of the training column).
   - We are picking the `median` to replace the missing values with because median is robust to outliers, and is safer than mean for skewed distributions
   - To create the processor object for this step to put in the `Pipeline` creation above, use sklearn `SimpleImputer` class. Make sure you pass the correct strategy mentioned above while constructing this object.
#### Step 3.2 `Scaling Numeric variables` --
   - Transforms numeric columns using: $x_\text{scaled} = \frac{x-\mu}{\sigma}$, where $\mu$ is the mean of the column, and $\sigma$ is the standard deviation of the column.
   - After scaling a column, mean of the column will become 0, and standard deviation will be 1.
   - Many machine learning models are very sensitive to the magnitude of the variables. So, scaling is important.
   - Please note: for binary variables, we do not need to scale because the values possess fixed scale. Scaling will destroy its meaning.
   - To create the processor object for this particular step to put in to the `Pipeline` creation above, use sklearn `StandardScaler` class.

#### Step 3.0 Construct a `Pipeline` object named `numeric_transformer`
* A sklearn Pipeline allows you to sequentially apply a list of transformers to preprocess the data. 
* * For this `numeric_transformer` you are going to apply the following two transformations: i) Missing value imputation step, ii) Standard Scaling step.
```python
from sklearn.pipeline import Pipeline
numeric_transformer = Pipeline(steps=[
    #list of steps as a list of tuples.
    #(name_of_step1, processor_object_1),
    #(name_of_step2, processor_object_2),
    # ...
])
```

#### Step 3.1 `Missing Value Imputation of the Numeric features` --
   - Replaces missing numeric values with a meaningful value (e.g., the median of the training column).
   - We are picking the `median` to replace the missing values with because median is robust to outliers, and is safer than mean for skewed distributions
   - To create the processor object for this step to put in the `Pipeline` creation above, use sklearn `SimpleImputer` class. Make sure you pass the correct strategy mentioned above while constructing this object.
#### Step 3.2 `Scaling Numeric variables` --
   - Transforms numeric columns using: $x_\text{scaled} = \frac{x-\mu}{\sigma}$, where $\mu$ is the mean of the column, and $\sigma$ is the standard deviation of the column.
   - After scaling a column, mean of the column will become 0, and standard deviation will be 1.
   - Many machine learning models are very sensitive to the magnitude of the variables. So, scaling is important.
   - Please note: for binary variables, we do not need to scale because the values possess fixed scale. Scaling will destroy its meaning.
   - To create the processor object for this particular step to put in to the `Pipeline` creation above, use sklearn `StandardScaler` class.

```python
# Step 3: Numeric column transformer
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler
from sklearn.impute import SimpleImputer

numeric_transformer = Pipeline(
	steps=[] #list of tuples for the steps: (), (), (), ...
)



# YOUR CODE HERE
numeric_transformer = Pipeline(steps=[
    ('imputer', SimpleImputer(strategy='median')),
    ('scaler', StandardScaler())
])


#raise NotImplementedError("Task 3, Step 3 not completed.")
```

```python
#Let me test your code above
assert len(numeric_transformer.steps)==2, "numeric_transformer must contain exactly 2 steps"
_,obj1 = numeric_transformer.steps[0]
_,obj2 = numeric_transformer.steps[1]
assert type(obj1).__name__ == 'SimpleImputer', "First step of numeric_transformer must be done by SimpleImputer"
assert type(obj2).__name__ == 'StandardScaler', "Second step of numeric_transformer must be done by StandardScaler"
print(f'Task 3, Step 3 test complete. All good.')
```

Output:
```
Task 3, Step 3 test complete. All good.
```

### Step 4: Prepare `Ordinal` Columns Preprocessing Pipeline
#### Step 4.0 Construct a `Pipeline` object named `ordinal_transformer`
* A sklearn Pipeline allows you to sequentially apply a list of transformers to preprocess the data. 
* * For this `ordinal_transformer` you are going to apply the following two transformations: i) Missing value imputation step, ii) Encode the ordinal columns with `OrdinalEncoder`.
```python
from sklearn.pipeline import Pipeline
ordinal_transformer = Pipeline(steps=[
    #list of steps as a list of tuples.
    #(name_of_step1, processor_object_1),
    #(name_of_step2, processor_object_2),
    # ...
])
```
#### Step 4.1: Handling Missing values in ordinal columns
* Before encoding we are going to use `SimpleImputer(strategy="most_frequent")` to fill the missing values. Missing values are replaced with the most common value in training.
* To create the processor object for this step to put in the `Pipeline` creation above, use sklearn `SimpleImputer` class. Make sure you pass the correct strategy mentioned above while constructing this object.
#### Step 4.2: Encode the ordinal columns with `OrdinalEncoder`
* To create the processor object for this particular step to put in to the `Pipeline` creation above, use sklearn `OrdinalEncoder` class:

```python

OrdinalEncoder(
    categories=list(the_orders.values()), #passing list(s) from the_orders dictionary
                                # which you created in Task3/Step2 above.
                                # as list of lists: [ [v1, v2, v3], [w1, w2, w3, w4]]
    handle_unknown="use_encoded_value",
    unknown_value=-1
)

```
* This preserves the ordering information of the ordinal feature columns.
* `handle_unknown='use_encoded_value'` tells the encoder not to raise an error when it sees a new category not present during training time, but instead to use a specific integer you define with `unknown_value` field. We chose `unknown_value=-1` in the `OrdinalEncoder` will encode any unknown value with -1 instead of causing an error.

```python
# Step 4: Ordinal column transformer

from sklearn.preprocessing import OrdinalEncoder
from sklearn.impute import SimpleImputer
from sklearn.pipeline import Pipeline

ordinal_transformer = Pipeline(
	steps=[] #list of tuples for the steps: (), (), (), ...
)

#YOUR CODE GOES HERE
ordinal_transformer = Pipeline(steps=[
    ('imputer', SimpleImputer(strategy='most_frequent')),
    ('ordinal', OrdinalEncoder(categories=list(the_orders.values()), 
                               handle_unknown='use_encoded_value', 
                               unknown_value=-1))
])
```

```python
#Let me test your code above
assert len(ordinal_transformer.steps)==2, "ordinal_transformer must contain exactly 2 steps"
_,obj1 = ordinal_transformer.steps[0]
_,obj2 = ordinal_transformer.steps[1]
assert type(obj1).__name__ == 'SimpleImputer', "First step of ordinal_transformer must be done by SimpleImputer"
assert type(obj2).__name__ == 'OrdinalEncoder', "Second step of ordinal_transformer must be done by OrdinalEncoder"
print(f'Task 3, Step 4 test complete. All good.')
```

Output:
```
Task 3, Step 4 test complete. All good.
```

### Step 5: `Nominal` Feature preprocessing
#### Step 5.0 Construct a `Pipeline` object named `nominal_transformer`
* A sklearn Pipeline allows you to sequentially apply a list of transformers to preprocess the data. 
* For this `nominal_transformer` you are going to apply the following two transformations: i) Missing value imputation step, ii) Encode the nominal columns with `OneHotEncoder`.
```python
from sklearn.pipeline import Pipeline
nominal_transformer = Pipeline(steps=[
    #list of steps as a list of tuples.
    #(name_of_step1, processor_object_1),
    #(name_of_step2, processor_object_2),
    # ...
])
```
#### Step 5.1: Handling Missing values in nominal columns
* Before encoding we are going to use `SimpleImputer(strategy="most_frequent")` to fill the missing values. Missing values are replaced with the most common value in training.
* * To create the processor object for this step to put in the `Pipeline` creation above, use sklearn `SimpleImputer` class. Make sure you pass the correct strategy mentioned above while constructing this object.

#### Step 5.2: Encode with `OneHotEncoder`
* We will use `OneHotEncoder(handle_unknown="ignore")` to encode the nominal columns.
* Nominal categories have no natural order. For instance, take a look at values in the `HISPMOM` column:

```
C:Cubans, M:Mexicans,  N:No, O:Colombians P:Peruvians, S:Salvadorans, U:Guatemalans
```
if encoded as 

```
0:Cubans, 1:Mexicans,  2:No, 3:Colombians 4:Peruvians, 5:Salvadorans, 6:Guatemalans
```
The model may incorrectly assume:

```
Cubans < Mexicans < No < Colombians < Peruvians < Salvadorans < Guatemalans

or, 

C < M < N < O < P < S< U
```
That is mathematically incorrect.
Instead, one-hot encoding produces:

| HISPMOM_C | HISPMOM_M | HISPMOM_N | HISPMOM_O | HISPMOM_P | HISPMOM_S | HISPMOM_U |
| ----------- | ----------- | ------------ |------------ |------------ |------------ |------------ |
| 1           | 0           | 0            |0            |0            |0            |0            |

* As you can see, each category gets its own binary column.
* If test data contains a category not seen in training, the option `handle_unknown="ignore"` in the `OneHotEncoder` will encode all one-hot columns for the nominal variable become 0 instead of causing an error. This prevents the model from crashing.

```python
#Step 5: Nominal column transformer

from sklearn.preprocessing import OneHotEncoder
from sklearn.impute import SimpleImputer
from sklearn.pipeline import Pipeline

nominal_transformer = Pipeline(
	steps=[] #List of tuples of each step in this pipeline: (), (), ...
)

# YOUR CODE HERE
nominal_transformer = Pipeline(steps=[
    ('imputer', SimpleImputer(strategy='most_frequent')),
    ('onehot', OneHotEncoder(handle_unknown='ignore'))
])


```

```python
#Let me test your code above
assert len(nominal_transformer.steps)==2, "nominal_transformer must contain exactly 2 steps"
_,obj1 = nominal_transformer.steps[0]
_,obj2 = nominal_transformer.steps[1]
assert type(obj1).__name__ == 'SimpleImputer', "First step of nominal_transformer must be done by SimpleImputer"
assert type(obj2).__name__ == 'OneHotEncoder', "Second step of nominal_transformer must be done by OneHotEncoder"
print(f'Task 3, Step 5 test complete. All good.')
```

Output:
```
Task 3, Step 5 test complete. All good.
```

### Step 6: Preprocessing all other feature columns
#### Step 6.0 Construct a `Pipeline` object named `default_transformer`
* A sklearn Pipeline allows you to sequentially apply a list of transformers to preprocess the data. 
* For this `default_transformer` you are going to apply the following transformation: Missing value imputation step.
```python
from sklearn.pipeline import Pipeline
default_transformer = Pipeline(steps=[
    #list of steps as a list of tuples.
    #(name_of_step1, processor_object_1)
])
```
#### Step 6.1: Handling Missing values in all other columns
* Before encoding we are going to use `SimpleImputer(strategy="most_frequent")` to fill the missing values. Missing values are replaced with the most common value in training.
* * To create the processor object for this step to put in the `Pipeline` creation above, use sklearn `SimpleImputer` class. Make sure you pass the correct strategy mentioned above while constructing this object.

```python
#Step 5: Default column transformer

from sklearn.preprocessing import OneHotEncoder
from sklearn.impute import SimpleImputer
from sklearn.pipeline import Pipeline

default_transformer = Pipeline(
	steps=[] #List of tuples of each step in this pipeline: (), (), ...
)

# YOUR CODE HERE
default_transformer = Pipeline(steps=[
    ('imputer', SimpleImputer(strategy='most_frequent'))
])
```

```python
#Let me test your code above
assert len(default_transformer.steps)==1, "default_transformer must contain exactly 1 step"
_,obj1 = default_transformer.steps[0]

assert type(obj1).__name__ == 'SimpleImputer', "First step of default_transformer must be done by SimpleImputer"

print(f'Task 3, Step 6 test complete. All good.')
```

Output:
```
Task 3, Step 6 test complete. All good.
```

### Step 7: Finalize a `ColumnTransformer` combining the preprocessing steps 3, 4, 5, and 6
* The sklearn's `ColumnTransformer` is a powerful class that allows you to apply different preprocessing pipelines (also called the transformers) to specific columns of your dataset simultaneously. This is particularly useful when dealing with heterogeneous data (e.g., a mix of numerical, ordinal, categorical ) that require different types and number of of transformations.
* Construct a `ColumnTransformer` object named `preprocess`:

```python
from sklearn.compose import ColumnTransformer

preprocess = ColumnTransformer(
  transformers=[
    #List of (name, transformer, columns) tuples specifying the transformer pipeline objects to be applied to subsets of the data..
    #(name_of_pipeline1, pipeline1_object, [list_of_columns_to_apply]),
    #(name_of_pipeline2, pipeline2_object, [list_of_columns_to_apply]),
    # ...
])
```
* Make sure your `preprocess` column applies the following 3 transformation pipelines you constructed in Steps 3, 4 and 5:
  - Numeric pipeline to the numeric feature columns, 
  - Ordinal pipeline to the ordinal feature columns,
  - Nominal pipeline to the nominal feature columns,
  - Default pipeline to all other feature columns.

```python
#Step 7: Combining the preprocessing pipelines (Steps 3, 4, 5, and 6)

from sklearn.compose import ColumnTransformer
preprocess = ColumnTransformer(
    transformers=[
        # list of transformation piplines on specific columns
    ]
)

# YOUR CODE HERE
preprocess = ColumnTransformer(transformers=[
    ('num', numeric_transformer, numeric_features),
    ('ord', ordinal_transformer, ordinal_features),
    ('nom', nominal_transformer, nominal_features),
    ('def', default_transformer, all_other_features)
])


```

```python
#Let me test your code above
assert len(preprocess.transformers)==4, "preprocess must contain exactly 4 steps for steps 3, 4, 5 and 6 you defined above."
_,obj1,cols1 = preprocess.transformers[0]
_,obj2,cols2 = preprocess.transformers[1]
_,obj3,cols3 = preprocess.transformers[2]
_,obj4,cols4 = preprocess.transformers[3]

assert (type(obj1).__name__ == 'Pipeline') and (type(obj2).__name__ == 'Pipeline') and (type(obj3).__name__ == 'Pipeline') and (type(obj4).__name__ == 'Pipeline'), "All 4 transformation pipelines must be of type `Pipeline`"

#assert preprocess.remainder == 'passthrough', "Make sure a other columns which did not go through any of the three preprocessing piplines passed through as is"
```

### Step 8: Train/fit `preprocess` transformer using the training split
* Now that we have sketched the `preprocess` transformer denoting how each of the feature columns will be transformed, we need it to train (or fit) it using the training split we created before (i.e., `X_train`).
* The `ColumnTransformer` has `fit()` method that takes the data to learn from. The `fit()` call updates all internal parameters of the transformer needed to transform any dataset we present to it.


```python
print("MAGER" in X_train.columns)
print(X_train.columns.tolist())
```

Output:
```
False
['SEX', 'MARITAL', 'FAGE', 'GAINED', 'VISITS', 'MAGE', 'FEDUCATION', 'MEDUCATION', 'TOTALP', 'BDEAD', 'TERMS', 'LOUTCOME', 'WEEKS', 'RACEMOM', 'RACEDAD', 'HISPMOM', 'HISPDAD', 'CIGNUM', 'DRINKNUM', 'ANEMIA', 'CARDIAC', 'ACLUNG', 'DIABETES', 'HERPES', 'HYDRAM', 'HEMOGLOB', 'HYPERCH', 'HYPERPR', 'ECLAMP', 'CERVIX', 'PINFANT', 'PRETERM', 'RENAL', 'RHSEN', 'UTERINE']
```

```python
#Step 8: Let's fit the preprocess transformer using the training split
preprocess.fit(X_train)

# YOUR CODE HERE

# Nothing to do here as I completed it for you above.
# I called the fit() function and passed the training split (X_train) to it. Please observe carefully.
# The fit() call updates all internal parameters of the transformer needed to transform any dataset we present to it.


```

Output:
```
ColumnTransformer(transformers=[('num',
                                 Pipeline(steps=[('imputer',
                                                  SimpleImputer(strategy='median')),
                                                 ('scaler', StandardScaler())]),
                                 ['MAGE', 'FAGE', 'GAINED', 'VISITS', 'WEEKS',
                                  'TOTALP', 'CIGNUM', 'DRINKNUM']),
                                ('ord',
                                 Pipeline(steps=[('imputer',
                                                  SimpleImputer(strategy='most_frequent')),
                                                 ('ordinal',
                                                  OrdinalEncoder(categories=[[0,
                                                                              1,
                                                                              2,
                                                                              3,
        
```

#### Step 9: Transform the training split with the fitted/trained transformer
* We now have a ready to use fitted `preprocess` transformer to apply the preprocessing steps we outlined on any dataset we present to it.
* The `ColumnTransformer` has `transform()` method that transforms the given data separately by each transformer, concatenate results and return it.
* Make sure to receive the result in `X_train_transformed` variable. We will give this transformed data to our models.

```python
#Step 9: Let's apply the transformer on the training split
X_train_transformed = preprocess.transform(X_train)

# YOUR CODE HERE

# Nothing to do here as I completed it for you above.
# I called the transform() function and passed the training split (X_train) to it. Please observe carefully.
# The transform() call applies all transformer pipelines, concatenate results and return it. So, I caught the result in the `X_train_transformed` variable.


```

#### Step 10: Transform the test split with the fitted/trained transformer
* Repeat step 8 by applying the same transformer onto the test split you created before.
* Make sure to receive the result in `X_test_transformed` variable. 

```python
#Step 10: Let's apply the same transformer on the test split
X_test_transformed = X_test

# YOUR CODE HERE
X_test_transformed = preprocess.transform(X_test)

```

```python
#Let me test your code above
assert type(X_test_transformed) == np.ndarray, 'Expected type mismatch. Preprocessing transformer returns a numpy array.'
```

## TASK 4: Linear Regression with Closed-form solution
Given the (`X_train_transformed`, `y_train`) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **Task 4.1**: Complete the **linear_regression_closed_form_training()** function:
  * It fits a linear regression model using the `closed-form solution` to obtain the coefficients, beta's (i.e., $\beta$'s), as a numpy array of `m+1` values (Please recall class lecture), where $m$ is the number of variables in `X_train_transformed` (the first argument to the function). 
  * The function returns the solution `beta`s (i.e., the `m+1` $\beta$ values).
* **Task 4.2**: Complete the **linear_regression_closed_form_predict()** function:
  * It takes a list of `m+1` beta values (i.e., `beta`s returned from the corresponding training function, and `X_test_transformed` (containing test samples each having `m` input variables). 
  * Then, using the provided beta values, predict each of the test samples, and let's name your predictions `y_pred` ($\hat{y}$). Return `y_pred` ($\hat{y}$) from the function.
* **Task 4.3**: Complete the **RMSE()** function: 
  * It takes two lists: `y_test` ($y$), `y_pred` ($\hat{y}$), where the first list represents ground truth (i.e., actual) target values for the given samples, and the second list represents a corresponding predicted values for exactly same number of samples in `y_test`. 
  * Compute and return the `Root Mean Squared Error (RMSE)` of the prediction:

$$
\text{RMSE}(\hat{y}, y) = \sqrt{\frac{1}{n} \sum_{i=1}^{n} (y_i - \hat{y}_i)^2}
$$

* In tasks 4.1, 4.2, and 4.3, PLEASE DO NOT USE ANY LIBRARY FUNCTION (e.g., scikit-learn) THAT DOES THE LINEAR REGRESSION FOR YOU or RMSE calculation.


```python
# Task 4.1: Complete the following function

# YOUR CODE HERE
def linear_regression_closed_form_training(X_train, y_train):
    # Add a column of ones for the intercept term (beta_0) 
    X_b = np.c_[np.ones((X_train.shape[0], 1)), X_train]
    
    # Calculate X^T * X 
    XTX = X_b.T @ X_b
    
    # Add a tiny value to the diagonal to prevent Singular Matrix error 
    # This is a standard way to handle collinearity in datasets 
    XTX_stable = XTX + 1e-5 * np.eye(XTX.shape[0])
    
    # Normal Equation: beta = (X^T * X)^-1 * X^T * y [cite: 383, 387]
    betas = inv(XTX_stable) @ X_b.T @ y_train
    return betas

    
    
```

```python
# YOUR CODE HERE
def linear_regression_closed_form_predict(betas, X_test):
    # Add intercept column to test data
    X_b = np.c_[np.ones((X_test.shape[0], 1)), X_test]
    # Prediction: y = X * beta
    y_pred = X_b @ betas
    return y_pred
```

<!-- grade_id: cell-task-10-ans-1, points: 0, grade: false, solution: true -->
```python
#Task 4.3: Complete the following function
# YOUR CODE HERE
def RMSE(y_test, y_pred):
    # Convert inputs to numpy arrays to handle list inputs from tests
    y_test_arr = np.array(y_test)
    y_pred_arr = np.array(y_pred)
    
    # Compute Root Mean Squared Error
    rmse_val = np.sqrt(np.mean((y_test_arr - y_pred_arr)**2))
    return rmse_val

```

* **Task 4.4**:
  * Call `linear_regression_closed_form_training()` function with `X_train_transformed`, `y_train` as arguments.
  * Store the returned results in variable, `betas_closed_form`.
  * Print `betas_closed_form`.
  

<!-- grade_id: cell-task-10-ans-2, points: 0, grade: false, solution: true -->
```python
#Task 4.4: 
betas_closed_form = []
#Your task here, call appropriate function to get these two
# YOUR CODE HERE
betas_closed_form = linear_regression_closed_form_training(X_train_transformed, y_train)


print(f'{betas_closed_form}')
```

Output:
```
[-5.17761325e+02  5.72364829e-02 -4.99596473e-05  1.73447225e-01
  4.01507468e-02  7.11625483e-01  7.10149041e-02 -1.07094819e-01
 -7.67777418e-04 -1.88022572e+02 -1.65389738e+02  3.77891426e+01
  3.77938446e+01  3.75269507e+01  3.77025453e+01  3.77975208e+01
  3.76604712e+01  3.79058323e+01  3.76065167e+01  3.76447129e+01
  2.62825705e+01  2.63541465e+01  2.62779825e+01  2.62961890e+01
  2.60694751e+01  2.61683753e+01  2.62661453e+01  2.59808353e+01
  2.60434183e+01  2.63708327e+01  4.13037206e+01  4.14864107e+01
  4.14265305e+01  4.13891338e+01  4.12765396e+01  4.14093522e+01
  4.14183468e+01  8.16233568e+01  8.17234564e+01  8.17829616e+01
  8.16233326e+01  8.16749949e+01  8.17143630e+01  8.15597575e+01
 -1.59622136e+01 -1.62201304e+01 -1.56649390e+01  6.46008181e-01
  6.07679541e-01  4.92570816e-01 -9.09524361e-02 -1.22768112e-01
 -6.77092406e-02 -9.27739992e-03  6.59280703e-03 -1.05755345e-02
  2.89841967e-01  1.58492322e-02 -3.62179007e-01  2.07394119e-02
 -2.01831300e-01 -3.47844
```

* **Task 4.5:**
	* Call `linear_regression_closed_form_predict`() function with `betas_closed_form`,`X_test_transformed`. 
	* Store the returned result in variable, `y_pred`.
  

<!-- grade_id: cell-task-10-ans-3, points: 0, grade: false, solution: true -->
```python
#Task 4.5:
y_pred = []

#Your task here. call appropriate function to get get y prediction for the supplied test dataset scaled

# YOUR CODE HERE
y_pred = linear_regression_closed_form_predict(betas_closed_form, X_test_transformed) 


```

* **Task 4.6:**
  * Call `RMSE`() function with `y_test` and `y_pred`. 
  * Store the returned result in `rmse_closed_form` variable.
  * Print `rmse_closed_form`.

<!-- grade_id: cell-task-10-ans-4, points: 0, grade: false, solution: true -->
```python
#Task 4.6
rmse_closed_form = 0
#Your task below. Call the function to compute RMSE score of the y predictions
# YOUR CODE HERE
rmse_closed_form = RMSE(y_test, y_pred)
print(f'rmse_closed_form = {rmse_closed_form}')


#raise NotImplementedError('Task 4.6 not complete.')

```

Output:
```
rmse_closed_form = 1.0347026650272304
```

<!-- grade_id: cell-task-10-test, points: 10, grade: true, solution: false -->
```python
#Let's check your Task 4 answers
assert abs(RMSE([1,2,3,4,5],[1,1,1,1,1])-2.449489742783178)<1e-4
assert abs(RMSE([100, 200, 300, 400, 500], [90, 190, 290, 390, 490])-10.0)<1e-4
assert y_pred.shape[0]==25350
```

## TASK 5: Linear Regression with Full-batch Gradient Descent
Given the (`X_train_transformed`, `y_train`) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **Task 5.1**: Complete the **linear_regression_full_batch_gd_training()** function:
  * It fits a linear regression model using the `full-batch gradient descent` to obtain the coefficients, beta's (i.e., $\beta$'s), as a numpy array of `m+1` values (Please recall class lecture), where $m$ is the number of variables in `X_train_transformed` (the first argument to the function). Make sure you use average of gradient in beta update formulas. Please use the `alpha` (i.e, the learning rate) and `nEpoch` (number of epochs) parameters in your implementation of the gradient descent algorithm.
  * The function returns the solution `beta`s (i.e., the `m+1` $\beta$ values).
* **Task 5.2**: Complete the **linear_regression_full_batch_gd_predict()** function:
  * It takes a list of `m+1` beta values (i.e., `beta`s returned from the corresponding training function, and `X_test_transformed` (containing test samples each having `m` input variables). 
  * Then, using the provided beta values, predict each of the test samples, and let's name your predictions `y_pred` ($\hat{y}$). Return `y_pred` ($\hat{y}$) from the function.

* In tasks 5.1, and 5.2, PLEASE DO NOT USE ANY LIBRARY FUNCTION (e.g., scikit-learn) THAT DOES THE LINEAR REGRESSION FOR YOU.

```python
#Task 5.1: Complete the function
# YOUR CODE HERE
def linear_regression_full_batch_gd_training(X_train, y_train, alpha, nEpoch):
    random.seed(554433)
    m, n = X_train.shape
    # Add intercept column of ones
    X_b = np.c_[np.ones((m, 1)), X_train]
    # Initialize betas to zeros
    betas = np.zeros(n + 1)
    
    for epoch in range(nEpoch):
        # Calculate gradients using the average gradient formula
        # Gradient = (2/m) * X^T * (X*beta - y)
        gradients = (2/m) * X_b.T @ (X_b @ betas - y_train)
        # Update betas
        betas = betas - alpha * gradients
        
    return betas
```

```python
#Task 5.2: Complete the following function
# YOUR CODE HERE
def linear_regression_full_batch_gd_predict(betas, X_test):
    # Add intercept column
    X_b = np.c_[np.ones((X_test.shape[0], 1)), X_test]
    # Prediction: y_pred = X * beta
    y_pred = X_b @ betas
    return y_pred
```

<!-- grade_id: cell-task-11-ans-1, points: 0, grade: false, solution: true -->
* **Task 5.3**:
  * Call `linear_regression_full_batch_gd_training()` function with `X_train_transformed`, `y_train`. Also run the training with learning rate, `alpha=0.01`, number of epochs, `nEpoch=1000`.
  * Store the returned results in variable, `betas_full_batch`.
  * Print `betas_full_batch`.
  

<!-- grade_id: cell-task-11-ans-2, points: 0, grade: false, solution: true -->
```python
#Task 5.3
betas_full_batch = []

#your task below
# YOUR CODE HERE
# YOUR CODE HERE
betas_full_batch = linear_regression_full_batch_gd_training(X_train_transformed, y_train, 0.01, 1000)
print(f'betas_full_batch = {betas_full_batch}')



```

Output:
```
betas_full_batch = [ 1.33322308e+00  8.77587558e-02  5.89203200e-03  1.65664406e-01
  3.83951856e-02  7.26231206e-01  2.67867583e-02 -1.27118139e-01
 -3.22271054e-03 -1.33322308e+00 -1.33322308e+00  4.06861084e-02
  5.57648518e-01  2.89205971e-01  1.37164137e-01  2.15257436e-02
  4.10019570e-03  3.05094353e-03  1.69289853e-02  2.62912475e-01
  3.35780464e-02  5.49474069e-01  3.53466463e-01  1.32147667e-01
  1.97173789e-02  1.94280606e-03  2.29033249e-03  7.86713101e-03
  2.25792809e-01  6.94637347e-03  1.93910561e-02  4.49474213e-01
  4.66859863e-01  1.52175337e-02  8.11721286e-02  2.89915929e-01
  1.11923537e-02  2.15237881e-02  4.21899732e-01  4.98469855e-01
  8.68520316e-03  8.40282321e-02  2.89898259e-01  8.71800808e-03
  7.95160127e-01  5.37689783e-01  3.73166453e-04  5.85731239e-01
  3.83589459e-01  3.63902379e-01  1.42723588e-01 -1.84557433e-02
  3.16534423e-03  1.42391274e-02  1.61367233e-03 -2.22926928e-03
  1.47700271e-01  5.84279636e-03 -8.29520780e-02  2.14810261e-03
 -3.76
```

```python
print(f'betas_full_batch = {betas_full_batch}')
```

Output:
```
betas_full_batch = [ 1.33322308e+00  8.77587558e-02  5.89203200e-03  1.65664406e-01
  3.83951856e-02  7.26231206e-01  2.67867583e-02 -1.27118139e-01
 -3.22271054e-03 -1.33322308e+00 -1.33322308e+00  4.06861084e-02
  5.57648518e-01  2.89205971e-01  1.37164137e-01  2.15257436e-02
  4.10019570e-03  3.05094353e-03  1.69289853e-02  2.62912475e-01
  3.35780464e-02  5.49474069e-01  3.53466463e-01  1.32147667e-01
  1.97173789e-02  1.94280606e-03  2.29033249e-03  7.86713101e-03
  2.25792809e-01  6.94637347e-03  1.93910561e-02  4.49474213e-01
  4.66859863e-01  1.52175337e-02  8.11721286e-02  2.89915929e-01
  1.11923537e-02  2.15237881e-02  4.21899732e-01  4.98469855e-01
  8.68520316e-03  8.40282321e-02  2.89898259e-01  8.71800808e-03
  7.95160127e-01  5.37689783e-01  3.73166453e-04  5.85731239e-01
  3.83589459e-01  3.63902379e-01  1.42723588e-01 -1.84557433e-02
  3.16534423e-03  1.42391274e-02  1.61367233e-03 -2.22926928e-03
  1.47700271e-01  5.84279636e-03 -8.29520780e-02  2.14810261e-03
 -3.76
```

* **Task 5.4:**
	* Call `linear_regression_full_batch_gd_predict`() function with `betas_full_batch`,`X_test_transformed`. 
	* Store the returned result in variable, `y_pred`.
  

<!-- grade_id: cell-task-11-ans-3, points: 0, grade: false, solution: true -->
```python
#Task 5.4
y_pred = []

#your task below
# YOUR CODE HERE
y_pred = linear_regression_full_batch_gd_predict(betas_full_batch, X_test_transformed)


```

* **Task 5.5:**
  * Call `RMSE`() function with `y_test` and `y_pred`. 
  * Store the returned result in `rmse_full_batch` variable.
  * Print `rmse_full_batch`.

<!-- grade_id: cell-task-11-ans-4, points: 0, grade: false, solution: true -->
```python
#Task 5.5
rmse_full_batch = 0
# YOUR CODE HERE

rmse_full_batch = RMSE(y_test, y_pred)
print(f'rmse_full_batch = {rmse_full_batch}')


```

Output:
```
rmse_full_batch = 1.049702619172222
```

## TASK 6: Linear Regression with Stochastic Gradient Descent
Given the (`X_train_transformed`, `y_train`) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **Task 6.1**: Complete the **linear_regression_stochastic_gd_training()** function:
  * It fits a linear regression model using the `Stochastic gradient descent` to obtain the coefficients, beta's (i.e., $\beta$'s), as a numpy array of `m+1` values (Please recall class lecture), where $m$ is the number of variables in `X_train_transformed` (the first argument to the function). Make sure you use average of gradient in beta update formulas. Please use the `alpha` (i.e, the learning rate) and `nEpoch` (number of epochs) parameters in your implementation of the gradient descent algorithm.
  * The function returns the solution `beta`s (i.e., the `m+1` $\beta$ values).
* **Task 6.2**: Complete the **linear_regression_stochastic_gd_predict()** function:
  * It takes a list of `m+1` beta values (i.e., `beta`s returned from the corresponding training function, and `X_test_transformed` (containing test samples each having `m` input variables). 
  * Then, using the provided beta values, predict each of the test samples, and let's name your predictions `y_pred` ($\hat{y}$). Return `y_pred` ($\hat{y}$) from the function.

* In tasks 6.1, and 6.2, PLEASE DO NOT USE ANY LIBRARY FUNCTION (e.g., scikit-learn) THAT DOES THE LINEAR REGRESSION FOR YOU.


```python
# Task 6.1
def linear_regression_stochastic_gd_training(X_train, y_train, alpha, nEpoch):
    random.seed(554433)
    m, n = X_train.shape
    # Add intercept column
    X_b = np.c_[np.ones((m, 1)), X_train]
    betas = np.zeros(n + 1)
    
    for epoch in range(nEpoch):
        # Shuffling at the start of each epoch for better convergence
        indices = np.random.permutation(m)
        for i in indices:
            xi = X_b[i:i+1]
            yi = y_train.iloc[i:i+1]
            
            # Using the average-style gradient update for stability [cite: 505]
            # beta = beta - alpha * (2/m) * xi^T * (xi*beta - yi)
            prediction = xi @ betas
            gradient = (2.0 / m) * xi.T @ (prediction - yi)
            
            betas = betas - alpha * gradient.flatten()
            
    return betas

```

<!-- grade_id: cell-task-12-ans-1, points: 0, grade: false, solution: true -->
```python
# Task 6.2
# YOUR CODE HERE
def linear_regression_stochastic_gd_predict(betas, X_test):
    X_b = np.c_[np.ones((X_test.shape[0], 1)), X_test]
    y_pred = X_b @ betas
    return y_pred
    


```

* **Task 6.3**:
  * Call `linear_regression_stochastic_gd_training()` function with `X_train_transformed`, `y_train`. Also run the training with learning rate, `alpha=0.01`, number of epochs, `nEpoch=50`.
  * Store the returned results in variable, `betas_stochastic`.
  * Print `betas_stochastic`.
  

<!-- grade_id: cell-task-12-ans-2, points: 0, grade: false, solution: true -->
```python
#Task 6.3
betas_stochastic = []
cpu_time_stochastic = 0
#your code below
# YOUR CODE HERE
betas_stochastic = linear_regression_stochastic_gd_training(X_train_transformed, y_train, 0.01, 50)


```

```python
print(f'betas_stochastic = {betas_stochastic}')
```

Output:
```
betas_stochastic = [ 9.03539563e-01  8.59515020e-02  6.22643394e-02  9.60689218e-02
  6.27509868e-02  4.67286678e-01 -3.17698083e-02 -1.39952613e-01
 -1.10455890e-02 -9.03539563e-01 -9.03539563e-01  3.93325062e-03
  6.91304395e-01  1.40754314e-01  1.48941040e-02  2.20909249e-03
  4.92165770e-04  3.32927032e-04  2.22285206e-03  4.73964618e-02
  3.35973516e-03  6.79730800e-01  1.59645342e-01  1.43179992e-02
  2.04498403e-03  2.12808004e-04  2.50832697e-04  1.20197189e-03
  4.20561135e-02  7.18976049e-04  2.62320710e-03  1.63088019e-01
  6.68855560e-01  1.94421215e-03  1.13275525e-02  5.43151569e-02
  1.38585567e-03  2.98584581e-03  1.66472002e-01  6.65507239e-01
  1.32159397e-03  1.16132303e-02  5.43772898e-02  1.26236232e-03
  5.07106425e-01  3.96430892e-01  2.24502686e-06  5.04456510e-01
  1.08905507e-01  2.90177545e-01  1.01191461e+00  9.34693806e-03
  2.12002640e-01  2.93871339e-02  3.34379879e-03  7.90039808e-03
  3.45391732e-02  1.01976138e-02  5.87955225e-03  1.38142297e-03
  6.57
```

* **Task 6.4:**
	* Call `linear_regression_stochastic_gd_predict`() function with `betas_stochastic`,`X_test_transformed`. 
	* Store the returned result in variable, `y_pred`.
  

<!-- grade_id: cell-task-12-ans-3, points: 0, grade: false, solution: true -->
```python
#Task 6.4

y_pred = []
#your code below
# YOUR CODE
y_pred = linear_regression_stochastic_gd_predict(betas_stochastic, X_test_transformed)


```

* **Task 6.5:**
  * Call `RMSE`() function with `y_test` and `y_pred`. 
  * Store the returned result in `rmse_stochastic` variable.
  * Print `rmse_stochastic`.

<!-- grade_id: cell-task-12-ans-4, points: 0, grade: false, solution: true -->
```python
rmse_stochastic = 0
#your code below
# YOUR CODE HERE
rmse_stochastic = RMSE(y_test, y_pred)

```

<!-- grade_id: cell-task-12-test, points: 10, grade: true, solution: false -->
```python
print(f'rmse_stochastic = {rmse_stochastic}')
```

Output:
```
rmse_stochastic = 1.221291059616642
```

## Task 7: Linear Regression with Mini-batch Gradient Descent

Given the (`X_train_transformed`, `y_train`) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **Task 7.1**: Complete the **linear_regression_minibatch_gd_training()** function:
  * It fits a linear regression model using the `Mini-batch gradient descent` to obtain the coefficients, beta's (i.e., $\beta$'s), as a numpy array of `m+1` values (Please recall class lecture), where $m$ is the number of variables in `X_train_transformed` (the first argument to the function). Make sure you use average of gradient in beta update formulas. Please use the `alpha` (i.e, the learning rate) and `nEpoch` (number of epochs) parameters in your implementation of the gradient descent algorithm.
  * The function returns the solution `beta`s (i.e., the `m+1` $\beta$ values).
* **Task 7.2**: Complete the **linear_regression_minibatch_gd_predict()** function:
  * It takes a list of `m+1` beta values (i.e., `beta`s returned from the corresponding training function, and `X_test_transformed` (containing test samples each having `m` input variables). 
  * Then, using the provided beta values, predict each of the test samples, and let's name your predictions `y_pred` ($\hat{y}$). Return `y_pred` ($\hat{y}$) from the function.

* In tasks 7.1, and 7.2, PLEASE DO NOT USE ANY LIBRARY FUNCTION (e.g., scikit-learn) THAT DOES THE LINEAR REGRESSION FOR YOU.


```python
# Task 7.1: Complete the following function
def linear_regression_minibatch_gd_training(X_train, y_train, alpha, nEpoch, batch_size):
    random.seed(554433)
    m, n = X_train.shape
    X_b = np.c_[np.ones((m, 1)), X_train]
    betas = np.zeros(n + 1)
    
    for epoch in range(nEpoch):
        # Shuffle indices to ensure random batches
        shuffled_indices = np.random.permutation(m)
        X_b_s = X_b[shuffled_indices]
        y_s = y_train.iloc[shuffled_indices]
        
        for i in range(0, m, batch_size):
            xi = X_b_s[i:i+batch_size]
            yi = y_s[i:i+batch_size]
            
            # Gradient update using the mini-batch average
            gradients = (2/len(xi)) * xi.T @ (xi @ betas - yi)
            betas = betas - alpha * gradients
            
    return betas
```

<!-- grade_id: cell-task-13-ans-1, points: 0, grade: false, solution: true -->
```python
# Task 7.2
def linear_regression_minibatch_gd_predict(betas, X_test):
    X_b = np.c_[np.ones((X_test.shape[0], 1)), X_test]
    y_pred = X_b @ betas
    return y_pred


```

* **Task 7.3**:
  * Call `linear_regression_minibatch_gd_training()` function with `X_train_transformed`, `y_train`. Also run the training with learning rate, `alpha=0.01`, number of epochs, `nEpoch=50`, batch size, `batch_size=32`.
  * Store the returned results in variable, `betas_minibatch`.
  * Print `betas_minibatch`.
  

<!-- grade_id: cell-task-13-ans-2, points: 0, grade: false, solution: true -->
```python
#Task 7.3
betas_minibatch = []
#your task below

# YOUR CODE HERE
betas_minibatch = linear_regression_minibatch_gd_training(X_train_transformed, y_train, 0.01, 50, 32) 

```

```python
print(f'betas_minibatch = {betas_minibatch}')
```

Output:
```
betas_minibatch = [ 1.64031151  0.07468891 -0.01093203  0.17075262  0.01109243  0.6805185
  0.10791098 -0.13170671 -0.06303402 -1.64031151 -1.64031151  0.30376851
  0.27290843  0.0600855   0.20787314  0.23193658  0.11776758  0.17680754
  0.11423     0.15493423  0.25220807  0.30061982  0.27357969  0.2657441
  0.12869018  0.06417553  0.10289992 -0.0247733   0.02204247  0.25512503
  0.15577692  0.36334278  0.26733646  0.23770979  0.12618733  0.24300228
  0.24695594  0.18235637  0.30165585  0.33938982  0.1786179   0.23842711
  0.26818001  0.13168444  0.9225662   0.68018486  0.03756044  0.61030517
  0.55665439  0.47335195 -0.0876673  -0.13283446 -0.08606507  0.00709305
  0.01009481 -0.00832684  0.29905517  0.03543959 -0.36283562  0.02482752
 -0.204314   -0.33021529 -0.65072074 -0.38829646  0.85410078 -0.47020372
 -0.08431671  0.03795553 -0.32430372]
```

* **Task 7.4:**
	* Call `linear_regression_minibatch_gd_predict`() function with `betas_minibatch`,`X_test_transformed`. 
	* Store the returned result in variable, `y_pred`.
  

<!-- grade_id: cell-task-13-ans-3, points: 0, grade: false, solution: true -->
```python
#Task 7.4
y_pred = []
#your ans below
# YOUR CODE HERE
y_pred = linear_regression_minibatch_gd_predict(betas_minibatch, X_test_transformed)

```

* **Task 7.5:**
  * Call `RMSE`() function with `y_test` and `y_pred`. 
  * Store the returned result in `rmse_minibatch` variable.
  * Print `rmse_minibatch`.

<!-- grade_id: cell-task-13-ans-4, points: 0, grade: false, solution: true -->
```python
#Task 7.5
rmse_minibatch = 0
#your ans below
# YOUR CODE HERE
rmse_minibatch = RMSE(y_test, y_pred)


```

<!-- grade_id: cell-task-13-test, points: 10, grade: true, solution: false -->
```python
print(f'rmse_minibatch = {rmse_minibatch}')
```

Output:
```
rmse_minibatch = 1.0384581679315203
```

## Task 8: Compare all 4 Linear Regression Models
* Given the 4 sets of results from the 4 experiments (from Tasks 4, 5, 6, 7) with closed form solution, full-batch gradient descent, stochastic gradient descent and mini-batch gradient descent, assign a string from the set {"closed-form", "batch-GD", "stochastic-GD", "minibatch-GD"} that demonstrated the best predictive performance in terms of RMSE to a variable named `best_model_name`.


<!-- grade_id: cell-task-14-ans, points: 0, grade: false, solution: true -->
```python
# Task 8
names = ["closed-form", "fullbatch-GD", "stochastic-GD", "minibatch-GD"]
rmsle_vals = [rmse_closed_form, 
              rmse_full_batch, 
              rmse_stochastic, 
              rmse_minibatch]

# YOUR CODE HERE
# Find the index of the minimum RMSE value
best_model_idx = np.argmin(rmsle_vals)

# Assign the corresponding name from the names list
best_model_name = names[best_model_idx]

# Be sure to delete the raise NotImplementedError line below


```

## Task 9: Hyper-parameter Tuning Exercise
* Given the (X_train_transformed, y_train) pairs denoting input matrix and output vector respectively, 
* call your implementation of Task 7, the minibatch gradient descent based linear regression for each of these learning rates: {0.001, 0.01, 0.05}, and batch sizes: {32, 64, 128, 256}
  - Please use nEpoch (number of epoch) = 50 in each experiment.

* For each of the linear regression model, using the computed beta values, predict the test samples provided in the "X_test_transformed" argument, and let's name your prediction "y_pred".
* Compute Root Mean Squared Error (RMSE) of your prediction using the RMSE() function you defined in Task 4.
* Finally, assign the learning rate that shows the best test performance to a variable called `best_alpha`, and also assign as a pandas dataframe named `summary` with 3 columns: {learning_rate, batch_size, test_RMSE} containing RMSE's of the 12 linear regression models, i.e., Cartesian product between the alphas and batch_sizes. Also, assign the best performing batch_size to a variable `best_batch_size`.
* Structure of the `summary` table (i.e., the Pandas dataframe):

| learning_rate | batch_size  | test_RMSE |
| --- | --- | --- |
| 0.001 | 32  | | |
| 0.001 | 64  | | |
| 0.001 | 128  | | |
| 0.001 | 256  | | |
| 0.01 | 32  | | |
| ... | ...  | | |



* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES ANY OF THE LINEAR REGRESSIONs.


<!-- grade_id: cell-task-16-ans, points: 0, grade: false, solution: true -->
```python
#Task 9
alphas = [0.001, 0.01, 0.05]
batch_sizes = [32, 64, 128, 256]
summary = pd.DataFrame(columns=['learning_rate', 'batch_size', 'test_RMSE'])
nIteration=1000
nEpoch=50

best_alpha = -1
best_beta = []
best_batch_size = -1

#your code below
# YOUR CODE HERE
results_list = []
min_rmse = float('inf')
for a in [0.001, 0.01, 0.05]:
    for b_size in [32, 64, 128, 256]:
        b = linear_regression_minibatch_gd_training(X_train_transformed, y_train, a, 50, b_size)
        y_p = linear_regression_minibatch_gd_predict(b, X_test_transformed)
        err = RMSE(y_test, y_p)
        results_list.append({'learning_rate': a, 'batch_size': b_size, 'test_RMSE': err})
        if err < min_rmse:
            min_rmse, best_alpha, best_batch_size, best_beta = err, a, b_size, b

summary = pd.DataFrame(results_list)







print(f"Summary = {summary}")

print(f"Best learning rate {best_alpha}")
print(f"Best batch size {best_batch_size}")
print(f'beta values: = {best_beta}')





```

Output:
```
Summary =     learning_rate  batch_size  test_RMSE
0           0.001          32   1.035628
1           0.001          64   1.036444
2           0.001         128   1.038590
3           0.001         256   1.043013
4           0.010          32   1.037343
5           0.010          64   1.035908
6           0.010         128   1.034802
7           0.010         256   1.037196
8           0.050          32   5.165098
9           0.050          64   1.046596
10          0.050         128   1.045929
11          0.050         256   1.058239
Best learning rate 0.01
Best batch size 128
beta values: = [ 1.60479233e+00  5.96258213e-02  3.43532146e-04  1.73102028e-01
  4.76626309e-02  7.08077714e-01  6.48703681e-02 -1.19773939e-01
 -7.15389080e-03 -1.60479233e+00 -1.60479233e+00  3.09068985e-01
  3.40582051e-01  7.51774560e-02  2.56608186e-01  1.94185968e-01
  5.70225836e-02  6.03927322e-02  1.20588925e-01  1.91165444e-01
  2.54229570e-01  3.54511053e-01  2.81481737e-01  2.98088601e-01
  1.6190
```

## Task 10:
* Utilizing the best trained linear regression model (so far), predict the target for each of the samples in the judge_dataset (`judge-without-labels-2.csv` loaded earlier).
    * I believe you will not forget to do the following before call in the prediction algorithm:
        - Save the ID values of the `judge dataset` into `ID_judge` and drop it from the judge dataframe.
        - Apply the trained `preprocess` pipeline on the `judge_dataset`.
    * Now, call the appropriate prediction function of the chosen model to obtain `y_pred`.
    * Prepare a pandas dataframe called `result` having columns: {ID, BWEIGHT}, where ID will the ID of the judge sample, and BWEIGHT is the corresponding y_pred value from your model prediction.
    * Save the dataframe as `my_submission.csv`.
* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES THE LINEAR REGRESSION.


<!-- grade_id: cell-task-17-ans-1, points: 0, grade: false, solution: true -->
```python
#Task 10
#your task below
#i) store the id in a separate variable and drop it from judge dataset
#ii) apply preprocess pipeline onto the dataset.

#(i) done below.
ID_judge = judge_dataset[['ID']]
judge_df = judge_dataset.drop(['ID'], axis=1)


judge_dataset_transformed = []

# YOUR CODE HERE
judge_dataset_transformed = preprocess.transform(judge_df)

#(ii) you need to do it correctly.


```

<!-- grade_id: cell-task-17-ans-2, points: 0, grade: false, solution: true -->
```python
#Task 10 (predict):
#your task is to predict using your best regression betas, and save it to `y_pred` variable
y_pred = []

# YOUR CODE HERE
# 1. Add the bias column (column of ones) to the transformed judge data
X_judge_b = np.c_[np.ones((judge_dataset_transformed.shape[0], 1)), judge_dataset_transformed]

# 2. Multiply the data by your best betas to get predictions
y_pred = X_judge_b @ best_beta

```

<!-- grade_id: cell-task-17-ans-3, points: 0, grade: false, solution: true -->
```python
#Task 10 (my_submission)
#your task below is to prepare a dataframe `result` with columns: {ID,BWEIGHT} as defined in task-10
result = pd.DataFrame()

# YOUR CODE HERE
# YOUR CODE HERE
# Create a DataFrame with the required ID and BWEIGHT columns
result = pd.DataFrame({
    'ID': ID_judge['ID'].values.flatten(),
    'BWEIGHT': y_pred
})






```

### End of Homework 2. Thanks for your time and effort.